# Vídeos que fueron tendencia en YouTube

### Disponible en Kaggle en:
https://www.kaggle.com/datasnaek/youtube-new


YouTube, el sitio web para compartir videos de fama mundial, mantiene una lista de los mejores videos de tendencias en la plataforma. Según la revista Variety, *Para determinar los videos más populares del año, YouTube utiliza una combinación de factores que incluyen la medición de las interacciones de los usuarios (número de visitas, compartidos, comentarios y me gusta). No son necesariamente los vídeos más vistos del año en general*. Los que se sitúan en la parte más alta de la lista de tendencias de YouTube suelen ser o bien vídeos musicales (como el famoso "Gagnam Style"), actuaciones de *celibrities* y / o reality shows, y vídeos virales variados de una persona aleatoria, cámara en mano.

Este conjunto de datos es un registro diario de los videos más populares de YouTube. Incluye varios meses de datos en vídeos de tendencias diarias de YouTube. Se incluyen datos para las regiones de EEUU, Reino Unido, Dinamarca, Canadá y Francia con hasta 200 videos de tendencias listados por día. Aquí solo usaremos los de EEUU, Canadá y Reino Unido. Los datos de cada país están en un archivo separado. Las variables incluyen el título del video, título del canal, tiempo de publicación, etiquetas, vistas, me gusta y no me gusta, descripción y recuento de comentarios. Se recopiló utilizando la API de YouTube.

### Variables y significado

Las variables utilizadas para describir cada vídeo son:

* video_id (string) ID único. Se ha asignado un video en la plataforma de YouTube. 
* trending_date (string) La fecha en que el video era tendencia 
* title (string) Título del vídeo
* channel_title (string) Título del canal de publicación en la categoría de plataforma
* category_id (string) El tipo de categoría del vídeo
* publish_time (string) La fecha de publicación del vídeo 
* tags (string) Etiquetas asociadas al vídeo, separadas por |
* views (entero) Número total de vistas en el vídeo.
* likes (entero) Número de Me gusta en el vídeo
* dislikes (entero) Número de No me gusta en el vídeo
* comment_count (entero) Un recuento total de comentarios en el video 
* comments_disabled (booleano) Si los comentarios estaban desactivados (true) o activados (false) en el video 
* ratings_disabled (booleano) Si la opción de dar me gusta o no al video está deshabilitada (true), en cuyo caso el número de  Me gusta y de No me gusta será 0. 
* video_error_or_removed (booleano) Si el video tiene algún error o se eliminó después de cargar el país
* description (string): Descripción textual

**Nombre completo del alumno:**  Pablo Benayas Penas

**INSTRUCCIONES**: en cada celda debes responder a la pregunta formulada, asegurándote de que el resultado queda guardado en la(s) variable(s) que por defecto vienen inicializadas a `None`. No se necesita usar variables intermedias, pero puedes hacerlo siempre que el resultado final del cálculo quede guardado exactamente en la variable que venía inicializada a None (debes reemplazar None por la secuencia de transformaciones necesarias, pero nunca cambiar el nombre de esa variable). **No olvides borrar la línea *raise NotImplementedError()* de cada celda cuando hayas completado la solución de esa celda y quieras probarla**.

Después de cada celda evaluable verás una celda con código. Ejecútala (no modifiques su código) y te dirá si tu solución es correcta o no. En caso de ser correcta, se ejecutará correctamente y no mostrará nada, pero si no lo es mostrará un error. Además de esas pruebas, se realizarán algunas más (ocultas) a la hora de puntuar el ejercicio, pero evaluar dicha celda es un indicador bastante fiable acerca de si realmente has implementado la solución correcta o no. Asegúrate de que, al menos, todas las celdas indican que el código es correcto antes de enviar el notebook terminado.

### Sobre los datasets youtube_USvideos.csv, youtube_CAvideos, youtube_GBvideos.csv se pide:

* Leer cada uno de estos ficheros en una variable, **sin intentar** que Spark infiera el tipo de dato de cada columna
* Puesto que existen columnas que contienen una coma enmedio del valor, en esos casos los valores vienen entre comillas dobles. Spark ya contempla esta posibilidad y puede leerlas adecuadamente **si al leer le indicamos las siguientes opciones adicionales** además de las que ya sueles usar: `.option("quote", "\"").option("escape", "\"")`.
* Asegúrate de que las **filas que no tienen el formato correcto sean descartadas**, indicando también la opción `mode` con el valor `DROPMALFORMED` como vimos en clase.
* Encadenadas con la operación de lectura `.csv()` de cada fichero, añadir para cada uno de los ficheros:
  - Una transformación para crear una nueva columna `pais` con una constante de tipo string (utiliza la función `F.lit("valor")`) indicando el país, que debe tomar como valores `"EEUU"`, `"CA"` y `"GB"` para EEUU, Canadá y Gran Bretaña respectivamente. Dicho valor será igual para todas las filas de cada uno de los tres DF, pero distinto de un DF a otro.
  * Una transformación que elimine la columna `description`, la cual contiene un texto que no analizaremos.
* Crear finalmente un nuevo DF `videosRawDF` en el que se hayan unido los tres DF anteriores. No debe ser cacheado todavía puesto que vamos a realizar operaciones de limpieza más adelante, y trabajaremos a partir de entonces con otro DF resultante.

<!-- grade_id: read_csv, points: 0, grade: false, solution: true -->
```python
# LÍNEA EVALUABLE, NO RENOMBRAR LAS VARIABLES
import pyspark.sql.functions as F
import numpy as np
from pyspark.sql.types import *
from pyspark.sql import Window
from functools import reduce


USvideosDF = spark.read\
                  .option("inferSchema", "false")\
                  .option("quote", "\"")\
                  .option("escape", "\"")\
                  .option("header", "true")\
                  .csv('gs://clase1_ntic/data/youtube_USvideos.csv', mode = "DROPMALFORMED")\
                  .withColumn('pais', F.lit('EEUU'))\
                  .drop('description') 

CAvideosDF = spark.read\
                  .option("inferSchema", "false")\
                  .option("quote", "\"")\
                  .option("escape", "\"")\
                  .option("header", "true")\
                  .csv('gs://clase1_ntic/data/youtube_CAvideos.csv', mode = "DROPMALFORMED")\
                  .withColumn('pais', F.lit('CA'))\
                  .drop('description') 

GBvideosDF = spark.read\
                  .option("inferSchema", "false")\
                  .option("quote", "\"")\
                  .option("escape", "\"")\
                  .option("header", "true")\
                  .csv('gs://clase1_ntic/data/youtube_GBvideos.csv', mode = "DROPMALFORMED")\
                  .withColumn('pais', F.lit('GB'))\
                  .drop('description') 

videosRawDF = USvideosDF.union(CAvideosDF).union(GBvideosDF) 
# YOUR CODE HERE
# raise NotImplementedError()
```

<!-- grade_id: read_csv_test, points: 1, grade: true, solution: false -->
```python
from pyspark.sql.types import DoubleType
assert((videosRawDF.count() == 136992) | (videosRawDF.count() == 120750))
assert("description" not in videosRawDF.columns)
assert((videosRawDF.where("pais = 'EEUU'").count() == 48137) | (videosRawDF.where("pais = 'EEUU'").count() == 40953))
assert((videosRawDF.where("pais = 'CA'").count() == 45560) | (videosRawDF.where("pais = 'CA'").count() == 40881))
assert((videosRawDF.where("pais = 'GB'").count() == 43295) | (videosRawDF.where("pais = 'GB'").count() == 38916))
```

* Las columnas `trending_date` y `publish_time` son en realidad de tipo fecha y de tipo timestamp (instante de tiempo), respectivamente, que Spark debería procesar como tales. Por otro lado, `likes`, `dislikes`, `comment_count` son de tipo entero, y `comments_disabled` es de tipo booleano. La columna `category_id` también lo es pero la utilizaremos más adelante para reemplazarla por sus verdaderas categorías, por lo que no vamos a modificarla ahora y la mantenemos como string. Partiendo de `vidosRawDF`, **reemplaza** cada una de estas columnas por su versión convertida al tipo de dato correcto en cada caso, utilizando `withColumn` con el mismo nombre de la columna existente. Para las columnas de tipo fecha y timestamp, el nuevo valor de la columna viene dado por el siguiente código:

        F.to_date("colName", "yy.dd.MM") # para fechas
        F.from_unixtime(F.unix_timestamp('colName', "yyyy-MM-dd'T'HH:mm:ss")).cast("timestamp") # para timestamp

* Después de las conversiones, eliminar todas las filas que tengan algún valor nulo
* El DF resultante de todas estas operaciones debe quedar almacenado en la variable `videosDF`, **cacheado**.


<!-- grade_id: convert_timestamp, points: 0, grade: false, solution: true -->
```python
# No olvides los imports que necesites...
# LÍNEAS EVALUABLES, NO RENOMBRAR LAS VARIABLES
videosDF = videosRawDF.withColumn('trending_date', F.to_date("trending_date", "yy.dd.MM"))\
                      .withColumn('publish_time', F.from_unixtime(F.unix_timestamp('publish_time',\
                                                   "yyyy-MM-dd'T'HH:mm:ss")).cast("timestamp"))\
                      .withColumn('likes', F.col('likes').cast(IntegerType()))\
                      .withColumn('dislikes', F.col('dislikes').cast(IntegerType()))\
                      .withColumn('comment_count', F.col('comment_count').cast(IntegerType()))\
                      .withColumn('comments_disabled', F.col('comments_disabled')\
                                                       .cast(BooleanType()))\
                      .dropna()\
                      .cache() 
# YOUR CODE HERE
# raise NotImplementedError()
```

<!-- grade_id: convert_timestamp_tests, points: 1, grade: true, solution: false -->
```python
dtypes = dict(videosDF.dtypes)
assert((videosDF.count() == 120739) | (videosDF.count() == 120746))
assert(dtypes["publish_time"] == "timestamp")
assert(dtypes["trending_date"] == "date")
assert(dtypes["likes"] == "int")
assert(dtypes["dislikes"] == "int")
assert(dtypes["comment_count"] == "int")
assert(dtypes["comments_disabled"] == "boolean")
assert(videosDF.is_cached)
```

Partiendo de `videosDF`:
* Traduce la categoría a un string utilizando la equivalencia indicada en el diccionario de la celda siguiente. Para ello puedes usar directamente la función `videosDF.replace(to_replace=diccionario, subset=['category_id'])` **que se aplica al DF completo** y devuelve un DF completamente nuevo en el que ha realizado en las columnas especificadas en la lista `subset` los reemplazamientos que le hayamos indicado en el diccionario `to_replace`. Por tanto, esta función **no se utiliza en combinación con `withColumn`** sino que debe ir fuera de la secuencia de transformaciones encadenadas, por ejemplo al principio de todo, y su resultado debe almacenarse en la variable `replacedCategoryDF`.

A continuación, partiendo de `replacedCategoryDF`: 
* Añade una nueva columna llamada `dias_hasta_viral` que contenga el número de días que han pasado entre la fecha en que se publicó un vídeo y el instante en el que se hizo viral. Para ello, utiliza `withColumn` en combinación con la función `F.datediff("columnaFuturo", "columnaPasado")`
* Añade otra nueva columna llamada `diasemana` que contenga el día de la semana en el que se ha publicado cada vídeo. Puedes usar la función `F.dayofweek("colName")` en combinación con `withColumn`.
* Vamos a empezar a utilizar la columna `tags`, para lo que necesitamos hacerla "usable". Encadenaremos estas tres transformaciones:
  - El primer paso consiste en convertir todas las palabras en minúsculas, para que exista más coincidencia entre etiquetas y podamos ver que la misma aparece en varios vídeos. Puedes usar la función `F.lower(F.col("columnName"))` en combinación con `withColumn` sobre la columna original. Atención: la función `lower` no admite directament el nombre de columna sino un objeto columna.
  - El segundo paso será eliminar el carácter `"` que rodea a la mayoría de términos individuales, puesto que no lo necesitamos ya que el separador entre términos es `|`. Puedes usar la función `F.regexp_replace("columnName", "\"", "")` en combinación con `withColumn` sobre la salida del apartado anterior, o incluso simplemente envolver a `F.lower(...)` con `regexp_replace`: `F.regexp_replace(F.lower(...), "\"", "")`.
  - El tercer paso consiste en convertirla en una columna de tipo **vector** de palabras, dividiendo la cadena de texto por el separador `|` que debe ser escapado poniendo una barra `\` delante porque `|` se utiliza también en expresiones regulares. **Reemplaza** la columna `tags` por el resultado de aplicarle la función `F.split("colName", "\|")`. Dicha función debe ser utilizada dentro de `withColumn`.
* El resultado debe quedar guardado en la variable `videosExtraInfoDF`.

<!-- grade_id: aniade_tiempos, points: 0, grade: false, solution: true -->
```python
# LÍNEA EVALUABLE, NO RENOMBRAR VARIABLES
# imports......
diccionario = {
   "1": "Film & Animation",
   "2": "Autos & Vehicles",
   "10": "Music",
   "15": "Pets & Animals",
   "17": "Sports",
   "18": "Short Movies",
   "19": "Travel & Events",
   "20": "Gaming",
   "21": "Videoblogging",
   "22": "People & Blogs",
   "23": "Comedy",
   "24": "Entertainment",
   "25": "News & Politics",
   "26": "Howto & Style",
   "27": "Education",
   "28": "Science & Technology",
   "29": "Nonprofits & Activism",
   "30": "Movies",
   "31": "Anime/Animation",
   "32": "Action/Adventure",
   "33": "Classics",
   "34": "Comedy",
   "35": "Documentary",
   "36": "Drama",
   "37": "Family",
   "38": "Foreign",
   "39": "Horror",
   "40": "Sci-Fi/Fantasy",
   "41": "Thriller",
   "42": "Shorts",
   "43": "Shows",
   "44": "Trailers"
}

replacedCategoryDF = videosDF.replace(to_replace=diccionario, subset=['category_id']) 
videosExtraInfoDF = replacedCategoryDF\
                      .withColumn('dias_hasta_viral', F.datediff('trending_date','publish_time'))\
                      .withColumn('diasemana', F.dayofweek("publish_time"))\
                      .withColumn('tags', F.regexp_replace(F.\
                                          lower(F.col("tags")), "\"", ""))\
                      .withColumn('tags', F.split("tags", "\|")) 
# YOUR CODE HERE
# raise NotImplementedError()
```

<!-- grade_id: aniade_tiempos_test, points: 2, grade: true, solution: false -->
```python
dtypesExtraInfo = dict(videosExtraInfoDF.dtypes)
assert(dtypesExtraInfo["category_id"] == "string")
assert(dtypesExtraInfo["tags"] == "array<string>")
assert(dtypesExtraInfo["diasemana"] == "int")
assert(dtypesExtraInfo["dias_hasta_viral"] == "int")
assert(videosExtraInfoDF.where("category_id = 'Education'").distinct().count() == 3103)
r = videosExtraInfoDF.select("video_id", "category_id", "diasemana", "tags").where("video_id == 'YVfyYrEmzgM'").head()
assert(r.category_id == "Education")
assert(r.diasemana == 2)
assert(("ted" in r.tags) & ("ted-ed" in r.tags) & ("ted education" in r.tags))
```

Partiendo de `videosExtraInfoDF`:

* Añadir una nueva columna `dias_viral` que contenga el **número de días durante los cuales un vídeo ha sido tendencia**. El dataset original contiene en muchos casos varias filas para un mismo vídeo, variando en la columna relativa al día en que el vídeo es tendencia. Lo que se pide es que la nueva columna contenga, **para cada fila**, la **diferencia** en días entre la fecha **mínima y máxima** que ese vídeo ha sido tendencia. El valor en ambas columnas se repetirá para todas las filas de cada vídeo, pero será distinto entre vídeos con distinto identificador. Primero deben calcularse una columna para el mínimo y otra para el máximo por cada vídeo **mediante agregaciones en una ventana (sin ordenar). No se debe utilizar la operación JOIN**. Una vez calculadas ambas columnas, la columna `dias_viral` será el resultado de aplicarles la función `F.datediff` utilizada anteriormente, y tras ello, se deben eliminar las columnas de mínimo y máximo (si se hubieran llegado a crear) que ya no son útiles.

* Tras esto, eliminar las filas duplicadas, atendiendo exclusivamente a la columna `video_id`, puesto que la información de qué días fue tendencia ya la habremos resumido en la columna `dias_viral`, y no queremos que cada vídeo aparezca varias veces y pueda falsear los resúmenes que haremos justo después. El resultado de estos dos apartados debe almacenarse en la variable `videosDiasViralDF`, que es aconsejable cachear porque lo usaremos varias veces más tarde.

¿Son todas las categorías igual de viralizantes? ¿Además de la categoría, es distinto el comportamiento según el país?

A continuación, y partiendo de `videosDiasViralDF`:

* Calcular un nuevo DF llamado `diasViralCategoriaPaisDF` con tantas filas como **categorías** y con tantas columnas como **países** (es decir, 3 más la columna de categorías), que contenga en cada celda el número **medio** de días que un vídeo de cada categoría permanece siendo viral en ese país. PISTA: Utilizar la función **pivot**.

* Por último, calcular otro DF llamado `videosPorDiasemanaDF` con tantas filas como **categorías** y tantas columnas como **días de la semana**, de forma que cada celda contenga el **número de vídeos** que se han publicado en esa categoría ese día de la semana. Tras el cálculo, **renombrar las columnas** para que sus nombres sean "Lunes", "Martes", "Miércoles", "Jueves", "Viernes", "Sábado", "Domingo" en lugar de los números enteros del 1 al 7.

<!-- grade_id: categoria_pais_semana, points: 0, grade: false, solution: true -->
```python
# LÍNEA EVALUABLE, NO RENOMBRAR VARIABLES
# import.....
ventanaVideoId = Window().partitionBy("video_id") 
videosDiasViralDF = videosExtraInfoDF\
                    .withColumn("trending_date_max",F.max("trending_date").over(ventanaVideoId))\
                    .withColumn("trending_date_min",F.min("trending_date").over(ventanaVideoId))\
                    .withColumn("dias_viral",F.datediff('trending_date_max','trending_date_min'))\
                    .drop('trending_date_max','trending_date_min')\
                    .dropDuplicates(["video_id"])\
                    .cache() 

diasViralCategoriaPaisDF = videosDiasViralDF.groupBy("category_id")\
                            .pivot("pais").agg(F.mean("dias_viral"))  

videosPorDiasemanaDF = videosDiasViralDF.groupBy("category_id").pivot("diasemana").count()
oldColumns = videosPorDiasemanaDF.columns
newColumns = ["category_id", "Lunes", "Martes", "Miércoles", "Jueves", "Viernes", 
              "Sábado", "Domingo"]

videosPorDiasemanaDF = reduce(lambda data, i: data\
                              .withColumnRenamed(oldColumns[i], newColumns[i]), 
                              range(len(oldColumns)), videosPorDiasemanaDF) 


videosPorDiasemana = videosPorDiasemanaDF.sort("category_id").collect() 
# YOUR CODE HERE
# raise NotImplementedError()
```

```python
videosPorDiasemana[0]
```

Output:
```
Row(category_id='Autos & Vehicles', Lunes=35, Martes=45, Miércoles=60, Jueves=36, Viernes=83, Sábado=33, Domingo=18)
```

<!-- grade_id: categoria_pais_semana_test, points: 3, grade: true, solution: false -->
```python
assert(dict(videosDiasViralDF.dtypes)["dias_viral"] == "int")
assert((videosDiasViralDF.count() == 30313) | (videosDiasViralDF.count() == 30318))
diasViralCategoriaPais = diasViralCategoriaPaisDF.sort("category_id").collect()
assert(diasViralCategoriaPais[0].category_id == "Autos & Vehicles")
assert(diasViralCategoriaPais[0].CA - 0.41025 < 0.001)
assert(diasViralCategoriaPais[3].category_id == "Entertainment")
assert(diasViralCategoriaPais[3].CA - 0.80932 < 0.001)
assert("Lunes" in videosPorDiasemanaDF.columns)
videosPorDiasemana = videosPorDiasemanaDF.sort("category_id").collect()
assert(videosPorDiasemana[0].category_id == "Autos & Vehicles")
assert(videosPorDiasemana[0].Lunes == 35)
assert(videosDiasViralDF.where("video_id = '-H90GPnH1q8'").head().dias_viral == 3)
assert("Martes" in videosPorDiasemanaDF.columns)
assert("Miércoles" in videosPorDiasemanaDF.columns)
assert("Jueves" in videosPorDiasemanaDF.columns)
assert("Viernes" in videosPorDiasemanaDF.columns)
assert("Sábado" in videosPorDiasemanaDF.columns)
assert("Domingo" in videosPorDiasemanaDF.columns)
assert(videosPorDiasemana[0].Domingo == 18)
```

Partiendo de `videosDiasViralDF`, donde cada vídeo solo aparece una vez, añadir las siguientes columnas:

* Una columna entera `ocurrencias_music` con el **número de ocurrencias** de la cadena **"music"** como **subcadena de cualquiera de los tags**. No es necesario que el tag sea exactamente igual a "music" sino que lo contenga como subcadena (recordemos que el DF de partida ya tiene todos los tags pasados a minúscula). Para ello, implementa una **UDF** `subtag_music_UDF` que devuelva `IntegerType()` y que envuelva a una función convencional de python llamada `subcadena_en_vector(tags)`. Esta última debe recibir como argumento una lista de strings, y comprobar cuántos elementos del vector contienen como subcadena a la palabra "music". Prueba su funcionamiento con la lista de tags `["a life in music", "music for life", "bso", "hans zimmer"]` que debería devolver 2. Finalmente invoca a `subtag_music_UDF` dentro de `withColumn` sobre la columna `tags`.

* Una columna `ocurrencia_media` de números reales con el número **medio** de apariciones de la palabra "music" como subcadena de tags en los vídeos similares al de la fila actual, entendiendo similares como aquellos del **mismo país y la misma categoría**. No debe utilizarse JOIN sino funciones de ventana.

* Una nueva columna `diff_porcentaje` que indique, en tanto por ciento, en qué medida el número de aparticiones de "music" está por encima o por debajo de la media de los vídeos de su misma categoría y país. Debe calcularse mediante **operaciones aritméticas entre columnas**, sin usar la función `F.when`, como la diferencia entre el número de aparticiones en el vídeo actual menos las apariciones medias de su país y categoría, dividido entre este último y multiplicado por 100.

El resultado de todas estas transformaciones debe quedar almacenado en la variable `videosOcurrenciasMusicDF`.

<!-- grade_id: ventana, points: 0, grade: false, solution: true -->
```python
# LÍNEA EVALUABLE, NO RENOMBRAR VARIABLES
# imports necesarios..........
def subcadena_en_vector(tags):
    return (sum([1 for c in tags if "music" in c])) 
subtag_music_UDF = F.udf(subcadena_en_vector, IntegerType())
subtag_music = videosDiasViralDF.withColumn("ocurrencias_music", subtag_music_UDF(F.col("tags"))) 

ventanaCategoriaPais = Window().partitionBy("pais","category_id") 
videosOcurrenciasMusicDF = subtag_music.withColumn('ocurrencia_media',\
                                            F.mean('ocurrencias_music')\
                                            .over(ventanaCategoriaPais))\
                                        .withColumn('diff_porcentaje',\
   ((F.col('ocurrencias_music')-F.col('ocurrencia_media'))/F.col('ocurrencia_media'))*100) 
# YOUR CODE HERE
# raise NotImplementedError()
```

<!-- grade_id: ventana_test, points: 3, grade: true, solution: false -->
```python
assert(subcadena_en_vector(["a life in music", "music for life", "bso", "hans zimmer"]) == 2)
r = videosOcurrenciasMusicDF.where("video_id = 'uSVW0aJdn9o'").head()
assert(r.ocurrencias_music == 2)
assert(r.ocurrencia_media - 1.0172278778386845 < 0.0001)
assert(r.diff_porcentaje - 96.612779060816 < 0.0001)
```